Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Lance"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    # YOUR CODE HERE
    best_sum = float(-99999) 
    #Make sure this is negative value of a huge number
    #because it will have problems with negative 
    #integers in the arrays
    left = 0 
    right = 0 
    
    for i in range(len(A)+1): 
        for j in range(i,len(A)+1): 
            current_sum = float(sum(A[i:j]))
            if current_sum > best_sum:
                best_sum = current_sum
                left = i
                right = j-1 
   

    
    return (left, right, best_sum)

A = [-2, -3, 4, -1, -2, 1]

print(bruteforce_max_subarray(A))

B = [-2, -3, 4, -1, -2, 1, 8]

print(bruteforce_max_subarray(B))
```

Output:
```
(2, 2, 4.0)
(2, 6, 10.0)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
No. We do not have to entirely run the array again. Assuming that we store the values of the sums we have acquired, getting the new array would be easier since the subsequent sums are basically the sums of our newly acquired sub arrays plus a new integer added. Since indexes (2,2): where (i,j), is the subarray with the highest sum, we just compare it with the sub arrays between the third integer and the newly added integer. 

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. The A and B have different maximums. Adding 8 at as the new last element is now include for the maximum subarray. Subarray A's maximum us (2,2,4). Subarray B's maximum is (2,6,10). Therefore, the new maximum-subarray needs to include 8. 

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    # YOUR CODE HERE
    if max(x) == x[0]:
        return x[0]
    
    max_sum = mx
    list_mx = x[0:-1] 
    current_sum = x[-1]
    for j in range(len(list_mx)-1, -1, -1): 
        current_sum += x[j]
        if current_sum > max_sum:
            max_sum = current_sum 
    
    return (max_sum)
    
B = [-2, -3, 4, -1, -2, 1, 8]

incremental_max_subarray(B[:1], 0)


```

Output:
```
-2
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    # YOUR CODE HERE
    
    
    
    raise NotImplementedError()
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
It will be more effective since it will have a run time of O(n). The reason being is that, it will only reiterate the loop once and the incrementations will perform the loop n+1 times. Since we are storing values of the maximum sum, we do not have to iterate the entire process again. 

```python

```