Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **50 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Lance"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work - Analyzing elementary sorting algorithms



## Question 1 [time estimate: 10 minutes]
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
**Line 7** moves the index of the sorted list and has two main functions. First, it shifts the index of the sorted list backwards and this allows the while loop to assess all then numbers in the sorted list going backwards. The other function is that it allows us the insert the key value prior to the sorted number that was bigger than the key value since A[i + 1] became the new A[i]. 

## Question 2 [time estimate: 15 minutes]
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertion_sort(A):
    #INSERT CODE HERE 
    for j in range(1, len(A)):
        key = A[j] #artificial memory for exchange 
        i= j-1
        while i >= 0 and A[i]>key: #we exchange if we found an unordered pair
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key #switch key to continue search 
    return A

insertion_sort([1,5,3,2,4])
insertion_sort([1,-2,0,-1])

```

Output:
```
[-2, -1, 0, 1]
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertion_sort([0]) == [0])
assert(insertion_sort([-1,1]) == [-1,1])
assert(insertion_sort([1,-1]) == [-1,1])
assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
The way the psuedocode works is that it starts its indexing at 2 where it assumes that the starting index of the list is at 1. However, in python, the index starts at 0 and therefore we just need to adjust the range of the for loop from 2 to 1. It will now index numbers at the rest of the list and now skipping the second number of the list. 

## Question 3 [time estimate: 15 minutes]
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertion_sort`,  your new Python function, `bubble_sort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubble_sort(A):
    # YOUR CODE HERE
    for i in range(0, len(A)-1): #range of the remaining numbers to be sorted 
        for j in range(len(A)-1, i, -1): #loop for exchanging numbers based on the unsorted list 
            if A[j] < A[j-1]:
                key = A[j] #artficial memory for exchange 
                A[j] = A[j-1]
                A[j-1] = key
    return A
                
bubble_sort([5,4,3,2,1])

```

Output:
```
[1, 2, 3, 4, 5]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4 [time estimate: 10 minutes]
The following pseudocode corresponds to a third sorting algorithm covered in this lesson, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free by providing at least three test cases. `selection_sort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selection_sort(A):
    # YOUR CODE HERE
    n = len(A)
    for i in range(0, n-1):
        minidx = i #assume that the first number is the lowest
        for j in range(i+1, n):
            if A[j] < A[minidx]:
                minidx = j #if we find a lower number then we switch our minidx
        #Swap 
        RAM = A[i] 
        A[i] = A[minidx]
        A[minidx] = RAM 

    return A 

selection_sort([5,4,3,2,1])
            
```

Output:
```
[1, 2, 3, 4, 5]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**